```python
# Initialize Otter
import otter
grader = otter.Notebook("lab.ipynb")
```

Output:
```
/Users/raunaksaluja/miniforge3/envs/dsc80/lib/python3.12/site-packages/otter/export/__init__.py:6: UserWarning: pkg_resources is deprecated as an API. See https://setuptools.pypa.io/en/latest/pkg_resources.html. The pkg_resources package is slated for removal as early as 2025-11-30. Refrain from using this package or pin to Setuptools<81.
  import pkg_resources
```

```python
import os
import pandas as pd
import numpy as np
np.set_printoptions(legacy='1.21')
import requests
import bs4
import lxml
import plotly.express as px
```

```python
recipes = pd.read_csv("RAW_recipes.csv")
interactions = pd.read_csv("interactions.csv")
```

```python
# Replace rating = 0 with NaN
interactions['rating'] = interactions['rating'].replace(0, np.nan)
interactions['rating'].value_counts(dropna=False)
recipes['submitted'] = pd.to_datetime(recipes['submitted'])
interactions['date'] = pd.to_datetime(interactions['date'])
```

```python
df = recipes.merge(
    interactions,
    left_on='id',
    right_on='recipe_id',
    how='left'
)

print("Merged shape:", df.shape)
df.head()
```

Output:
```
Merged shape: (234429, 17)
                                   name      id  minutes  contributor_id  \
0  1 brownies in the world    best ever  333281       40          985201   
1    1 in canada chocolate chip cookies  453467       45         1848091   
2                412 broccoli casserole  306168       40           50969   
3                412 broccoli casserole  306168       40           50969   
4                412 broccoli casserole  306168       40           50969   

   submitted                                               tags  \
0 2008-10-27  ['60-minutes-or-less', 'time-to-make', 'course...   
1 2011-04-11  ['60-minutes-or-less', 'time-to-make', 'cuisin...   
2 2008-05-30  ['60-minutes-or-less', 'time-to-make', 'course...   
3 2008-05-30  ['60-minutes-or-less', 'time-to-make', 'course...   
4 2008-05-30  ['60-minutes-or-less', 'time-to-make', 'course...   

                                      nutrition  n_steps  \
0      [138.4, 10.0, 50.0, 3.0, 3.0, 19.0, 6.0]      
```

```python
avg_rating = (
    df.groupby('id')['rating']
      .mean()
      .rename('average_rating')
)

df = df.merge(avg_rating, on='id', how='left')
```

```python
# Convert string to list
nutrition_cols = [
    'calories', 'total_fat_PDV', 'sugar_PDV',
    'sodium_PDV', 'protein_PDV',
    'saturated_fat_PDV', 'carbs_PDV'
]

df[nutrition_cols] = (
    df['nutrition']
    .str.strip("[]")
    .str.split(",", expand=True)
    .astype(float)
)
df['log_minutes'] = np.log1p(df['minutes'])
```

```python
# basic health indicator based on tags (edit keywords if you want)
healthy_keywords = [
    "healthy", "low-carb", "low carb", "low-fat", "low fat",
    "low calorie", "low-calorie", "high protein", "high-protein",
    "low sugar", "low-sugar"
]

tags_str = df["tags"].astype(str).str.lower()
df["is_healthy_tag"] = tags_str.apply(lambda s: any(k in s for k in healthy_keywords))

# useful transforms
df["log_minutes"] = np.log1p(df["minutes"])
df["log_calories"] = np.log1p(df["calories"])

# choose target column for plots
# if you have average_rating, use it; otherwise use rating
target_col = "average_rating" if "average_rating" in df.columns else "rating"
```

```python
fig1 = px.histogram(
    df.dropna(subset=["calories"]),
    x="log_calories",
    nbins=60,
    title="Distribution of Recipe Calories (log1p transformed)",
    labels={"log_calories": "log(1 + calories)"}
)
fig1
```

```python
fig2 = px.histogram(
    df.dropna(subset=["minutes"]),
    x="log_minutes",
    nbins=60,
    title="Distribution of Recipe Cooking Time (log1p transformed)",
    labels={"log_minutes": "log(1 + minutes)"}
)
fig2
```

```python
df.shape
```

Output:
```
(234429, 28)
```

```python
df[['average_rating','minutes','n_steps','n_ingredients','calories','is_healthy_tag']].describe()
```

Output:
```
average_rating       minutes        n_steps  n_ingredients  \
count   231652.000000  2.344290e+05  234429.000000  234429.000000   
mean         4.676531  1.067897e+02      10.017835       9.071506   
std          0.497757  3.285975e+03       6.442265       3.823039   
min          1.000000  0.000000e+00       1.000000       1.000000   
25%          4.500000  2.000000e+01       6.000000       6.000000   
50%          4.862745  3.500000e+01       9.000000       9.000000   
75%          5.000000  6.000000e+01      13.000000      11.000000   
max          5.000000  1.051200e+06     100.000000      37.000000   

           calories  
count  234429.00000  
mean      419.52887  
std       583.22359  
min         0.00000  
25%       170.70000  
50%       301.10000  
75%       491.10000  
max     45609.00000
```

```python
df['is_healthy_tag'].value_counts()
```

Output:
```
is_healthy_tag
False    153284
True      81145
Name: count, dtype: int64
```

```python
healthy = df[df["is_healthy_tag"]]["average_rating"].dropna()
not_healthy = df[~df["is_healthy_tag"]]["average_rating"].dropna()

obs_stat = healthy.mean() - not_healthy.mean()

obs_stat
```

Output:
```
-0.0033591948955473683
```

```python
n_perm = 3000
diffs = []

data = df[["average_rating","is_healthy_tag"]].dropna()

for i in range(n_perm):

    shuffled = data.copy()
    shuffled["is_healthy_tag"] = np.random.permutation(shuffled["is_healthy_tag"])

    g1 = shuffled[shuffled["is_healthy_tag"]]["average_rating"]
    g2 = shuffled[~shuffled["is_healthy_tag"]]["average_rating"]

    diff = g1.mean() - g2.mean()

    diffs.append(diff)

diffs = np.array(diffs)
```

```python
p_value = np.mean(diffs <= obs_stat)

print("Observed difference:", obs_stat)
print("p-value:", p_value)
```

```python
import plotly.express as px

fig = px.histogram(
    diffs,
    nbins=50,
    title="Permutation Distribution of Rating Difference"
)

fig.add_vline(x=obs_stat, line_color="red")

fig.show()
```

```python

```

Tests for q1 (points: None)
```
>>> os.path.exists('lab06_1.html')
True
```
```
>>> q01_html = bs4.BeautifulSoup(open("lab06_1.html", encoding='utf-8'), features='lxml')
>>> len(q01_html.find_all('image')) == 0
True
```
failure_message: you should not use the <image> tag to add images


Tests for q2 (points: None)
```
>>> extract_book_out[1] == extract_book_url
True
```
```
>>> isinstance(get_product_info_out, dict)
True
```
```
>>> 'Category' in get_product_info_out.keys()
True
```
```
>>> get_product_info_out['Rating'] == 'Two'
True
```
```
>>> scrape_books_out.shape == (1, 11)
True
```
```
>>> scrape_books_out['Rating'][0] == 'Four'
True
```
```
>>> scrape_books_out['Title'][0] == 'Sharp Objects'
True
```


Tests for q3 (points: None)
```
>>> history.shape == (19, 10)
True
```
```
>>> history.date.iloc[-1] == '2019-02-01'
True
```
```
>>> (len(stats[0]), len(stats[1])) == (6, 6)
True
```
```
>>> (float(stats[0][1:-1]) < 10) == True
True
```
```
>>> float(stats[1][:-1]) > 1 == True
True
```
```
>>> stats[1][-1] == 'B'
True
```


Tests for q4 (points: None)
```
>>> comments.shape == (18, 5)
True
```
```
>>> comments.loc[5, 'by'] == 'RobAtticus'
True
```
```
>>> comments.loc[5, 'time'].day == 31
True
```
